Substrate evaporation drives collective construction in termites

Author Response

The following is the authors’ response to the previous reviews.

This paper now provides a convincing presentation of valuable results of the drivers of nest construction for one termite species, and they briefly discuss possible relevance to other termite species. However, the authors have not yet addressed how their results may be important outside the field of termite nest construction. I could imagine the significance of the paper being elevated to important if there is a broader discussion about the impact of this work, e.g., the relevance of the results, the approach, and/or next steps to related fields outside of termite nest construction.

Reading our manuscript again, we have to agree with the reviewer that we mostly focused the discussion of our results in the context of termite construction, without attempting to generalise to other systems. To some extent we still defend this choice, as we prefer not to make too many claims on the relevance of our results beyond what we can reasonably support with our own experimental results. However, we thought that it would be appropriate – as suggested by the reviewer – to add at least one paragraph to indicate how our results could be extrapolated to other systems. This new paragraph is now at the end of the discussion section.

Here we elaborate a bit further on this point: first of all, while termites certainly build the most complex structures found in the natural world, there aren’t many other animals that are capable of collectively building complex structures. Typically, collective building activity is limited to highly social (typically eusocial) animals, but other social insects, such as ants and wasps, are phylogenetically distant from termites, their nests are often different (the large majority of ant nests only comprise excavated galleries with little construction, while wasp nests tend to comprise multiple repeated patterns that could be produced from stereotyped individual behaviour). Because of these differences, drawing a comparison between the mechanisms that regulate termite architecture and those that regulate other forms of animal architecture would be too speculative. One domain, however, where similar mechanisms to those that we describe here could operate is that of pattern formation at the cellular and tissue level, where surface curvature was shown to drive different phenomena from cell migration to tissue growth. A comment on this is now added in the manuscript at the very end of the discussion.

Similarly, on a related note, as someone not directly in the field of termite nest construction but wanting to understand the system (and the results) presented here in a broader context, I found the additional information about species and natural habitat very helpful and interesting, though I was rather disappointed to find it relegated to supplementary material where most readers will not see it.

We considered this suggestion to present more information about the natural nesting habits of the termites that we study into the main text, but eventually we decided to leave it as supplementary only. We feel that the nesting habits of the termites that we studied here are not too central to the problem that we want to focus on, of how they coordinate their building activity. In fact, there is a large variety of nesting habits across termite genera and species, but we believe that, at a basic level, the mechanisms that we describe here would also apply to species with different nesting habits, because our results are consistent with what is described in the scientific literature for other termite species. As our introduction is already a bit long, we left this description of Coptotermes nesting habits in the supplementary material, where, hopefully, it will still be accessible and useful to readers interested in finding this information.

When providing responses to reviewers, please directly address the reviewers’ comments point-by-point rather than summarizing comments and responding to summaries.

We apologize for our previous way to respond to comments and thanks the reviewer for his remark as we learn to navigate through the eLife reviewing system (where some comments are repeated in the overall assessment and in the feed-back of individual reviewers).

Figure 2 colors: Panels A and E and maybe B do not seem colorblind-friendly. I suggest modifying the colormaps to address this.

We have changed the colormaps of figures A,B and E which are now colorblind-friendly.

Line 180: This system is not in equilibrium. Perhaps the authors mean "steady-state?" I suggest reviewing language to ensure that the correct technical terms are used.

We have now corrected this.

Changes from the earlier version to the revised version

Abstract

Abstract

Termites build complex nests which are an impressive example of self-organization. We know that the coordinated actions involved in the construction of these nests by multiple individuals are primarily mediated by signals and cues embedded in the structure of the nest itself. However, to date there is still no scientific consensus about the nature of the stimuli that guide termite construction, and how they are sensed by termites. In order to address these questions, we studied the early building behavior of Coptotermes gestroi termites in artificial arenas, decorated with topographic cues to stimulate construction. Pellet collections were evenly distributed across the experimental setup, compatible with a collection mechanism that is not affected by local topography, but only by the distribution of termite occupancy (termites pick pellets at the positions where they are). Conversely, pellet depositions were concentrated at locations of high surface curvature and at the boundaries between different types of substrate. The single feature shared by all pellet deposition regions was that they correspond to local maxima in the evaporation flux. We can show analytically and we confirm experimentally that evaporation flux is directly proportional to the local curvature of nest surfaces. Taken together, our results indicate that surface curvature is sufficient to organize termite building activity, and that termites likely sense curvature indirectly through substrate evaporation. Our findings reconcile the apparently discordant results of previous studies.

Introduction

Termites are known for their ability to build some of the most complex nests and shelters found in nature (Hansell, 2005; Perna and Theraulaz, 2017). The construction of these structures is achieved through the collective actions of multiple individual workers (up to thousands or millions in large termite colonies) each performing the collection, transportation and deposition of elementary pellets. In order to produce functionally meaningful structures, it is essential that all these different workers operate in a coordinated, coherent way, each continuing the work started by their colony mates, rather than undoing it.

Termites rely on individual memory and proprioception to guide their behavior (see e.g. Bardunias and Su, 2009a), but these individual abilities are considered not sufficient to explain nest construction more generally. Instead, it is believed that building activity is largely guided by signals and cues embedded directly in the structure of the nest itself, through a regulation principle identified for the first time by Grassé, who named it stigmergy (Grassé, 1959; Camazine et al., 2001). In stigmergy-mediated nest-building, the probability for an individual insect to pick or to drop a pellet at a particular location is modulated by stimuli encountered at that location, such as the geometry of a growing pillar, or the presence of a chemical signal released by the queen or by other workers.

However, several years since Grassé’s early observations, there is still not a consensus on the exact nature of the stigmergic stimuli involved in regulating termite construction. Pheromones might be implicated in this regulation. Bruinsma (1979) found evidence for the role of a building pheromone released by the queen in the construction of the royal chamber of the termite Macrotermes subhyalinus. Computer simulation studies, aimed at reproducing the building behavior of termites and ants, also assume the existence of a “cement pheromone” added to the building material (Khuong et al., 2011, 2016; Heyde et al., 2021). In these simulation studies, the main and essential role of a cement pheromone is to allow initial pellet depositions to continue growing by differentiating them from regions of pellet collection, through differential pheromone marking. Experimental evidence in support for such cement pheromone in termite construction is weak: while individual workers can recognize freshly deposited nest material, they could simply be attracted to an unspecific colony odor while exhibiting the same behavioral patterns that they would exhibit also in the absence of chemical marking (Petersen et al., 2015). In other words, it is not clear if cement pheromones are required to drive termite building activities, or unspecific chemical cues would be sufficient, and it is also unclear if chemical stimuli modulate the building behavior of termites directly, or only indirectly, by affecting their density of presence.

Recent experimental studies by various authors have indicated that morphological and environmental features associated with some nest structures are strong stimuli that could by themselves guide termite construction activity. These include elevation (Fouquet et al., 2014), humidity gradients (Soar et al., 2019), and surface curvature (Calovi et al., 2019). The strong attractiveness of digging sites for termite aggregation means that in all these studies digging and deposition actions mostly come in pairs, which prevents us from identifying the genuine cues for pellet collection and deposition (Bardunias and Su, 2009b, 2010; Fouquet et al., 2014; Green et al., 2017). For example, in Calovi et al. (2019) termites are shown to preferentially aggregate in concave regions of a surface and they would simply rearrange nest material (both digging and building) at those locations. Even if digging sites provide a template for pellet deposition (Fouquet et al., 2014; Green et al., 2017), and for this reason digging and construction often co-localize in space, it is clear that building and digging cannot completely overlap, or the two activities would simply cancel one the effect of the other: termites must be able to differentiate between the sites of these contrasting activities through digging- or building-specific cues.

Some of the published computer simulation models of termite nest-building do not require a specific construction pheromone and assume instead that termites respond to cues naturally embedded in the nest structure itself. For example, the model proposed in Ocko et al. (2019) indicates that a generic “colony odor” undergoing advection and diffusion within the nest could provide a sufficient cue for determining the overall mound shape, so leaving a possible role of a construction pheromone only for the structuring of small scale nest features such as pillars and walls. Facchini et al. (2020) further proposed a model in which also small scale nest features can be produced in the absence of a construction pheromone, by assuming that termites respond to the local curvature of these emerging nest features. While these models reproduce a number of structures observed in real termite nests, the building rules implemented in the models are not empirically validated from direct observations of the building behavior of termite workers. As it stands, there is no conclusive evidence that the rules implemented in these models reflect the actual nest-building strategies of termites.

Here we aim to test whether geometric and physical cues embedded in the nest material are sufficient to explain termite construction. Specifically we want to disentangle how elevation, surface curvature, and substrate evaporation affect pellet deposition and collection. We do this by combining three different approaches. (i) We perform building experiments in which populations of termites are confronted with pre-existing building cues such as pillars, walls, and pre-made pellets of building material unmarked with pheromones. Using video-tracking, we monitor the presence of individual termites and we implement high throughput video-analysis to detect the time and location of individual pellet collection and deposition events. These experiments allow us to test the specific role played by each cue on stimulating pellet collection, pellet deposition, or termite aggregation. (ii) By running a computational model of nest building (Facchini et al., 2020) directly on the same structures that we provide to termites (obtained from 3D scans of our experimental setups), we can test exactly what building patterns we should expect under the simple assumption that termite depositions are driven by the local curvature of nest surfaces as the only construction cue. (iii) Finally, we develop a “chemical garden” experiment, on identical setups to those offered to termites, that allow us to visualize the sites of stronger water evaporation on the surface of the built structure. Overall, our approach allows us to demonstrate, both analytically and experimentally, the relation between deposition probability, surface curvature, and evaporation.

Results

Below, we report the observations of de novo building experiments performed with small experimental groups of Coptotermes gestroi termites confronted with a thin disk of humid clay covered with pre-made pellets unmarked with pheromones and decorated with pre-prepared clay features. In the first series of experiments the pre-prepared features were two pillars at the center of the clay disk as shown in figure 1.

Sketch of the experimental setup (left) and snapshot of one experiment (E66) before termites were added to the setup (right). The white marks on the picture give the scale of the setup, with the distance between successive marks being 1, 3, and 5 cm.

Pellet collection activity was distributed homogeneously all over the clay disk that we provided at the center of the experimental arena. Conversely, deposition activity was concentrated at the tips of pre-existing pillars, and along the edges of the clay disk itself. Figure 2A reports the heatmap of cumulative depositions P (D) and collections P (C) for one experiment (E66) with two pillars as topographic cues. A snapshot of the same experiment is reported in figure 3A. In figure 2B, we also report the same results for five experiments for which our analyses were most reliable because of the absence of spontaneous digging. Across all experiments, collections were widely distributed across the clay disk (i.e. where initial pellets are) while depositions were peaked at radii R ∼ 0.4 cm and R ∼ 2.5 cm which correspond to the top of pillars and to the edges of the experimental arena. Thus, termites do not show a preference for where they collect pellets while they target specific regions when depositing, which suggests that those regions must express a strong stimulus for deposition.

Top: (A) cumulative heatmaps of deposition (P (D); red) and collection activity (P (C); green) normalized by their respective mean values for one experiment (E66), colorbars are the same as in panel (E); (B) cumulative depositions (top) and collections (bottom) per unit area as a function of the Petri dish radius for experiments E58, E63, E65, E66, and E76, all histograms have been normalized and sum up to 1; (C) comparison among experimental depositions (in red), surface curvature (in blue) shown in Fig. 3C, and depositions predicted by simulations (black) shown in Fig. 3C, all the quantities are computed along the radial cut shown in panel (A), depositions are normalized by their maximum value and curvature is in mm−1; (D) cumulative occupancy heatmap normalized by its mean value for E66; (E) depositions (P (D|O); red) and collections (P (C|O); green) conditional to cumulative normalized occupancy for E66.

Top row: snapshots of a building experiment with “pillars” cue (E66) (A) and a building experiment with “wall” cue (E78, in B). Bottom row: snapshots of 3D simulations initiated with copies of the experimental setup E66 (C,D) and E78 (E,F) in which nest growth is entirely determined by the local surface curvature (based on our previously described model (Facchini et al., 2020)). Snapshots C and E refer to t=0, D and F refer to t=9 (dimensionless). The color map corresponds to the value of the mean curvature at the interface air-nest. White indicates convex regions and black indicates concave regions. The scale bars correspond to 1 cm.

To validate this hypothesis, we analyzed how building activity is related to the termite occupancy in the experimental setup. In figure 2D, we report the normalized cumulative occupancy of termites P (O) in the experimental setup. Occupancy is high close to the pillars and to the Petri dish walls, has intermediate values within the clay disk, and drops at the top of the pillars and right outside of the clay disk (i.e. precisely where deposits are recorded). To estimate how position and building activity are related, we report the conditional probabilities of depositing P (D|O) and collecting P (C|O) given termite occupancy. They are defined as the ratio between P (D) to P (O) and P (C) to P (O) as reported in figure 2E, and explained in section S.III of the Supplementary Information (SI). The probability P (D|O) reaches values 10 times larger than P (C|O) which confirms that our topographic cues and the clay disk edges specifically drive early building activity.

Earlier version: Focusing on topographic cues, we observe that pillar tips are the most curved part of the topography but also the most elevated one. In order to disentangle the respective roles of curvature and elevation in guiding pellet deposition, we considered a different setup where a thin wall replaced the two pillars in the center of the arena as shown in figure 3B. This way, the top edge of the wall is still a region of both high elevation and high surface curvature but elevation is constant every-where while curvature has local maxima at the tips. We report that the top edge attracted many deposits, but pellet deposition focused at the wall tips pointing to curvature, rather than elevation, as the dominant cue in 7 out of 11 experiments (SI table S1).

Revised version: Focusing on topographic cues, we observe that pillar tips are the most curved part of the topography but also the most elevated one. In order to disentangle the respective roles of curvature and elevation in guiding pellet deposition, we considered a different setup where a thin wall replaced the two pillars in the center of the arena as shown in figure 3B. This way, the top edge of the wall is still a region of both high elevation and high surface curvature but elevation is constant everywhere while curvature has local maxima at the tips. We report that the top edge attracted many deposits, but pellet deposition focused at the wall tips pointing to curvature, rather than elevation, as the dominant cue in 7 out of 11 experiments (SI table S1).

We wanted to further test to what extent the patterns that we observe are consistent with termites only responding to local substrate curvature, as opposed to responding also to other cues. To this end, we ran a model of nest construction that we have previously developed (Facchini et al., 2020) using 3D scans of the experimental arena – before the introduction of termites – as a starting template for the simulations (see Materials and Methods for simulation details). The simulation model implements one single construction rule which is a building response to local surface curvature and as such informs us about the possible building outcome that we could expect under the simplified assumption that construction is driven by surface curvature only, in the absence of any other cues. We should note that our model is phenomenological, and having one single parameter -that broadly defines the scale at which curvature is sensedis not intended to predict quantitative details such as the speed of the building process, or fine details of the shape of pillars and other structures. These simulations yielded the results shown in figure 3C-F. Experiments and simulations show a fair agreement as pellet depositions and initial growth concentrate in the same regions which are those where the surface is the most convex, as depicted in white in figures 3C-F. To make our comparison more readable in Fig. 2C we report a radial cut of the deposition heatmap (red), the surface curvature (blue) and the amount of depositions predicted by the simulations (black; see SI section S.VII for technical details of this analysis). The three curves show a good agreement and they all are peaked close to R = 4 mm which corresponds to the pillar tips. The qualitative agreement between experimental results and curvature-based simulations supports the idea that, at least on a first approximation, surface curvature alone is a sufficient cue that could guide termite depositions.

The edges of the clay disk were not included in the simulations because we could not characterize them properly with our scanning device (see SI, section S.III). However, in additional experiments with no topographic cues (SI, Fig. S4) most depositions happened precisely at the edges of the clay disk. It is possible that the very small edge of the clay disk provided a sufficient stimulus, in terms of local curvature, to elicit pellet depositions. However, the curvature cue was likely very weak at those locations as edges were smoothed out to gently match the surface of the Petri dish. Thus, we expect this region to bear a cue other than curvature (or elevation) which is also attractive for pellet depositions.

Earlier version: Trying to identify this additional building cue, we recall that the clay disk is maintained constantly humid. The edges of the disk arena mark then the limit between a humid region and the surrounding dry periphery. Also, the clay tone remained unchanged during experiments which suggests that moisture is constantly evaporating from the clay disk while being replenished in water from below, and that the overall process is stationary. To confirm this hypothesis, we measured the value of humidity and temperature both in the central and peripheral regions. We observed a net increase in humidity and a net decrease in temperature coming from outside to inside the clay disk which is the footprint of evaporation (SI, Fig. S5). Inside and outside the clay disk, both quantities remained relatively stable indicating that the system is roughly at equilibrium. Evaporation is a complex process, but close enough to the evaporating substrate, humidity transport happens by diffusion (Langmuir, 1918; Hisatake et al., 1993) and it is hence fully determined by the humidity gradient. In agreement with previous studies (Soar et al., 2019), we can show with scaling arguments that our termites live in such a viscous boundary layer (see SI, section S.IX). We can hence focus our attention on this specific region without loss of generality. For example, the humidity transition at the edges of the clay disk implies that the humidity gradient must be pronounced there and evaporation with it. In the diffusive regime, the evaporation flux is directly proportional to the surface curvature of the evaporating substrate (see SI, section S.VIII for a mathematical proof). As a demonstration for our topographic cues, we have computed the equilibrium solution for the humidity field h in a cubic volume bounded by pillars and wall experimental templates at the bottom. In the diffusive regime, this corresponds to solving the Laplace equation Δh = 0 while imposing a relative humidity h = 100% at the bottom boundary, and h = 70% at the top boundary (see S.VI in the SI for details) which was the average value of h in our experimental room. In figures 4A and 4B we have reported the contour plot of the magnitude of the humidity gradient |∇h| for this stationary solution. We can observe that the humidity gradient is maximum at the tip of pillars and at the lateral tips of the wall top edge, that are the most curved parts in the two different cues (see figures 3C and 3E for a direct comparison). As such, curvature and evaporation are two completely interchangeable stimuli everywhere except at the edge of the clay disk, where the transition between clay and perspex material corresponds to a strong humidity gradient in spite of weak surface curvature. We then propose that evaporation flux can explain by itself the deposition patterns observed in our experiments.

Revised version: Trying to identify this additional building cue, we recall that the clay disk is maintained constantly humid. The edges of the disk arena mark then the limit between a humid region and the surrounding dry periphery. Also, the clay tone remained unchanged during experiments which suggests that moisture is constantly evaporating from the clay disk while being replenished in water from below, and that the overall process is stationary. To confirm this hypothesis, we measured the value of humidity and temperature both in the central and peripheral regions. We observed a net increase in humidity and a net decrease in temperature coming from outside to inside the clay disk which is the footprint of evaporation (SI, Fig. S5). Inside and outside the clay disk, both quantities remained relatively stable indicating that the system is roughly in a steady state. Evaporation is a complex process, but close enough to the evaporating substrate, humidity transport happens by diffusion (Langmuir, 1918; Hisatake et al., 1993) and it is hence fully determined by the humidity gradient. In agreement with previous studies (Soar et al., 2019), we can show with scaling arguments that our termites live in such a viscous boundary layer (see SI, section S.IX). We can hence focus our attention on this specific region without loss of generality. For example, the humidity transition at the edges of the clay disk implies that the humidity gradient must be pronounced there and evaporation with it. In the diffusive regime, the evaporation flux is directly proportional to the surface curvature of the evaporating substrate (see SI, section S.VIII for a mathematical proof). As a demonstration for our topographic cues, we have computed the steady state solution for the humidity field h in a cubic volume bounded by pillars and wall experimental templates at the bottom. In the diffusive regime, this corresponds to solving the Laplace equation Δh = 0 while imposing a relative humidity h = 100% at the bottom boundary, and h = 70% at the top boundary (see S.VI in the SI for details) which was the average value of h in our experimental room. In figures 4A and 4B we have reported the contour plot of the magnitude of the humidity gradient |∇h| for this stationary solution. We can observe that the humidity gradient is maximum at the tip of pillars and at the lateral tips of the wall top edge, that are the most curved parts in the two different cues (see figures 3C and 3E for a direct comparison). As such, curvature and evaporation are two completely interchangeable stimuli everywhere except at the edge of the clay disk, where the transition between clay and perspex material corresponds to a strong humidity gradient in spite of weak surface curvature. We then propose that evaporation flux can explain by itself the deposition patterns observed in our experiments.

Top row: contour of the humidity gradient ∇h obtained solving the Laplace equation Δh = 0 in a cubic domain with a humid bottom boundary h = 100% (in brown) which is mapped from 3D scans of the experimental setup in E66 (A) and E78 (B). At the top boundary h is fixed to h = 70% which was the average value of humidity in our experimental room. Note that h is the relative humidity, thus the magnitude of the humidity gradient |∇h| is measured in mm−1, i.e |∇h| = 0.1 mm−1 means a humidity variation of 10% over 1 mm. Each contour corresponds to a variation of 0.015 mm−1. Pillar tips are associated with a strong humidity gradient; the top of the wall is also associated with a strong humidity gradient, although not as strong as at the pillar tips. Also note that humidity gradient at the top corners of the wall is 30% stronger than on the rest of the top edge. Bottom row: snapshots of chemical garden experiments initiated with “pillars” cue (C), and with “wall” cue (D). All the scale bars correspond to 1 cm.

To support this hypothesis we designed a chemical garden experiment that allows us to visualize the evaporation field in our setup. We prepared identical experimental setups as those used with termites, but this time we did not put any termites in the experimental arena. Instead we replaced the deionized water that was used to humidify the clay in the termite experiments with a saturated saline solution of water and NaHCO3. In this configuration, water evaporation is accompanied by the deposition of salt, which allows to build a chart of evaporation flux. Typical results are shown in figures 4C and 4D. Salt deposits appear in the form of white traces or bumpy deformations of the clay surface. Remarkably the distribution of the salt deposits matches very closely the regions of highest building activity by termites, both being more pronounced at the edge of the clay disk and at the top of topographic cues (pillar and walls). This result indicates that termite deposition probability covaries with the evaporation flux, which is consistent with our hypothesis of evaporation as the strongest cue for deposition. One may notice that salt traces are less pronounced on the wall (Fig. 4D) than on the pillars (Fig. 4C). This is consistent with the amplitude of the humidity gradient in a stationary diffusive regime as shown in figures 4B and 4A respectively: the maximum amplitude of the gradient is weaker in the wall case. Also, one observes that while having its maxima at the lateral tips (|∇h| ∼ 0.13 mm−1), the humidity gradient is strong all over the top edge of the wall (|∇h| ∼ 0.10 mm−1). Possibly, this could explain why in the case of wall cues there is a comparatively much larger variability in the patterns of termite construction (see SI Fig. S3). When termites are confronted with the ‘wall’ cue they are likely to start pellet depositions at the wall tips, but initial pellet depositions started at other locations on the top edge do occur. Whatever the initial choice, depositions are then preserved and reinforced by positive feed-back mechanisms which can lead to tightly selected deposition patterns also in the case of the ‘wall’ cue.

Note that, in the picture of depositions being attracted by evaporation flux, depositions observed at the edges of the clay disk, also agree with our previous growth model driven only by curvature (Facchini et al., 2020). In fact, the edge of the clay disk is almost flat (weak convexity) for a termite walking across, but it is also a thin cusp (high convexity) of humid material which is strongly evaporating, similarly to what happens at the edge of a liquid drop and causes the formation of well known coffee stains (Deegan et al., 1997). For a better comprehension, this apparent contrast is explained in the sketch of figure S7 (SI).

Discussion

Several experimental studies have tried to identify the cues that mediate termite construction, alternately indicating elevation (Fouquet et al., 2014), digging activity (Green et al., 2017), humidity transitions (Soar et al., 2019; Bardunias et al., 2020), or surface curvature (Calovi et al., 2019) as the relevant stimuli to drive pellet depositions. However, the fact that termites often concentrate their building activity in the immediate proximity of digging sites (Fouquet et al., 2014; Green et al., 2017; Bardunias and Su, 2009b, 2010) did not allow identifying which of these stimuli were specific digging and building cues, or simply generic cues for termite activity and aggregation. The cues themselves identified by different studies were different, leaving it unclear which, if any, were the relevant ones sensed by termites.

Here, by providing loose and unmarked pellets, we were able to prompt building activity without digging and to quantify collections and depositions as separate actions. We observed that all pellets are progressively displaced and that collections happened in a relatively random fashion. On the contrary, depositions concentrated at specific parts of the experimental arena which are the tips of the topographic cues and the edges of the clay disk. The conditional probability of deposition given termite occupancy is high there, indicating that those regions precisely drive termite depositions rather than generically attracting termite aggregation. The alternative use of pillars and walls as topographic cues allowed us to disentangle the role of elevation and curvature and pointed towards curvature as the most attractive stimulus for deposition. By simulating the building process with a model in which construction activity is driven by curvature only (Facchini et al., 2020), we obtained a good match with experimental results, indicating that curvature alone is sufficient to explain pellet depositions on topographic cues (pillars and walls).

Surface curvature is a powerful morphogenetic organizer for 3D structure formation as it can drive the formation of pillars, walls and convoluted surfaces (Facchini et al., 2020), all features that are observed in the nests of various termite species. Here, we are able to demonstrate a close coupling between surface curvature and the flux of evaporation from a surface, so providing a link to a possible stimulus sensed directly by termites. This also allows us to reconcile previous discordant results pointing alternately to curvature (Calovi et al., 2019; Facchini et al., 2020) or to humidity (Soar et al., 2019; Bardunias et al., 2020) as the relevant stimuli. The idea itself of a relation between curvature and evaporation is not new, as already a century ago, Langmuir (1918) showed that close enough to the surface of a water droplet, evaporation scales as the inverse of the radius (i.e. as the mean curvature) of the droplet (SI, section S.VIII.A). Our system is more complex than isolated spheres but our calculations in S.IX (SI) show that a relation between evaporation and curvature still holds at the termite scale.

As a further, direct, confirmation of our hypothesis, our chemical garden experiments clearly show that the correspondence between surface curvature and evaporation flux is relevant in our experimental setup.

Earlier version: It is well known that termites are particularly sensitive to the humidity of their environment, because their small size and soft cuticle put them in constant danger of desiccation. Coptotermes termites in particular are wetwood termites that can only survive in high-humidity environments such as moist wood or soil. For example, in laboratory experiments Zukowski and Su (2017), C. for-mosanus termites died within a few days when maintained at 72.9% or less relative humidity, but survived well when humidity was 98%. It is hence not surprising that they can sense and respond to humidity with their behavior, for example (Arab and Costa-Leonardo, 2005) tested the digging behavior of Coptotermes gestroi in wet sand with different levels of moisture and showed that tunnel length and the number of secondary branches changed when soil moisture increased from 5% to 15%. Coptotermes gestroi were also able to discriminate between chambers with different relative humidity and, after 12 hours, almost all termites were in the highest humidity chamber (98%), leaving the other chambers with 75% or less relative humidity empty (Gautam and Henderson, 2011). These results (which are similar also to other results testing termite response to chambers with different soil moisture) indicate that -given a sufficient amount of time-termites can detect a difference of humidity from 75% to 98% over a spatial scale of centimeters. Recent field and laboratory experiments have shown that humidity can affect termite behavior also in the context of building activity (Carey et al., 2019), by triggering nest expansion events (Bardunias et al., 2020; Carey et al., 2021). Even more interestingly, Soar et al. (2019) showed that moisture flux affects termite building activity (both in terms of pellet collection and deposition). Our experiments confirm this trend and suggest that moisture variations not only prompt or inhibit termite building activity, but constitute a local blueprint for construction.

Revised version: It is well known that termites are particularly sensitive to the humidity of their environment, because their small size and soft cuticle put them in constant danger of desiccation. Coptotermes termites in particular are wetwood termites that can only survive in high-humidity environments such as moist wood or soil. For example, in laboratory experiments Zukowski and Su (2017), C. formosanus termites died within a few days when maintained at 72.9% or less relative humidity, but survived well when humidity was 98%. It is hence not surprising that they can sense and respond to humidity with their behavior, for example (Arab and Costa-Leonardo, 2005) tested the digging behavior of Coptotermes gestroi in wet sand with different levels of moisture and showed that tunnel length and the number of secondary branches changed when soil moisture increased from 5% to 15%. Coptotermes gestroi were also able to discriminate between chambers with different relative humidity and, after 12 hours, almost all termites were in the highest humidity chamber (98%), leaving the other chambers with 75% or less relative humidity empty (Gautam and Henderson, 2011). These results (which are similar also to other results testing termite response to chambers with different soil moisture) indicate that -given a sufficient amount of timetermites can detect a difference of humidity from 75% to 98% over a spatial scale of centimeters. Recent field and laboratory experiments have shown that humidity can affect termite behavior also in the context of building activity (Carey et al., 2019), by triggering nest expansion events (Bardunias et al., 2020; Carey et al., 2021). Even more interestingly, Soar et al. (2019) showed that moisture flux affects termite building activity (both in terms of pellet collection and deposition). Our experiments confirm this trend and suggest that moisture variations not only prompt or inhibit termite building activity, but constitute a local blueprint for construction.

One may question if the humidity gradient in our experimental setup was strong enough for termites to sense it. Below, we use simple arguments to show that this condition was very likely fulfilled. Indeed, one can estimate that the value of the humidity gradient far from the disk edges and the topographic cues was |∇h|0 = δh/δ ∼ 0.15 mm−1, where δ = 2 mm is the thickness of the diffusive boundary layer (see S.IX, SI), and δh = 30% is the humidity variation between the clay disk surface and the experimental room.

Then, using the qualitative results of chemical gardens shown in 4C and 4D, one can conclude that |∇h|0 represents an estimation of the lower boundary of the humidity gradient experienced by termites in our experimental setup. Using diffusive simulations reported in figures 4A and 4B (see also S.VI SI), one can then quantify the relative importance of the humidity gradient at the tips of the topographic cues as 10 times larger, which gives an upper boundary |∇h|max ∼ 1 mm1.

As our termites are millimetre size, the estimated lower and upper boundary correspond to humidity variation of 10 to 100% across a distance not larger than their body length. We then expect that such variations must be sensed by termites, as they are larger than those that Coptotermes gestroi termites were shown to be able to discriminate over much larger spatial scales in the experiments reviewed above.

Our experiments do not support a role for a putative cement pheromone, added by termites to the building material, which would stimulate pellet depositions. In fact, construction occurred reliably on our provided building cues, even if they only comprised fresh clay and sterilized pellets with no pheromone markings. Our simulations further indicate that qualitatively similar construction results can be obtained without assuming a role for construction pheromone. We can hence exclude the influence of a cement pheromone, at least during the early choice of the deposition sites, in agreement with recent experiments by other authors (Fouquet et al., 2014; Petersen et al., 2015; Green et al., 2017). We should point to the fact, however, that in our experiments the building substrate was constantly moist throughout the entire duration of the experiments. It is possible that in some occurrences of nest building behavior, including in termites’ natural environment, moisture may not constantly replenish the porous wall of the growing structure. We suggest that under these conditions the evaporation flux is maintained by the humidity that is naturally embedded in recently dropped pellets, which makes the construction process self-sustainable and is consistent with the hypothesis of a viscous boundary layer extending with termite activity (Soar et al., 2019). In practice, it would be very hard to distinguish between such a scenario and one which involves a putative cement pheromone added directly to manipulated pellets by termites. More generally, while we do not rule out a possible role of pheromones in termite building behavior (mediating for instance termite aggregation), we have shown that pheromones are not necessary to explain the early deposition patterns that we see in our experiments.

In this study, we have focused on understanding how termites respond to well-controlled predefined stimuli. However, collective nest construction is a dynamic process and the deposition of new pellets by termites constantly changes the shape and the porosity of the evaporating substrate, potentially affecting nest growth through positive or negative feedback. Recent studies have shown that termites can control the size of the pellets used for nest construction, and indirectly also the porosity of nest walls (Zachariah et al., 2017). In turn, substrate porosity is known to play an important role for ventilation and drainage of the nest (Singh et al., 2019) and the moisture content of pellets can also affect the mechanical properties of the mound itself (Zachariah et al., 2020). In relation to our own experiments, however, our scaling analyses (SI S.IX) indicate that our conclusions are relatively robust to changes in substrate porosity and moisture content. For example, porosity only controls the time scale of water uptake from the reservoir by capillary rise, which must be small enough to keep the clay disk hydrated, and this assumption remains valid up to mm-size pores in the new construction. Similarly, for local curvature, the addition of new pellets to regions of high convexity is likely to make the surface less smooth than the initial topography, and such additional “roughness” can only increase the effect of focusing evaporation at those locations.

Previous work by Calovi and collaborators (2019) had pointed to an effect of surface curvature on termite construction behavior. While our two studies emphasize the same point, we should note that our results and the results reported in Calovi et al. (2019) are not entirely consistent, because in our experiments, pellet depositions are attracted by convex features, while in Calovi et al. (2019) termite activity was concentrated at regions of maximum concavity. As this previous study did not distinguish digging from deposition activity, we believe that their measure is a correlation between concavity and digging activity, not building. The fact that concave regions should attract digging activity is predicted by our model (SI, section S.V) and was visible also in our experiments where concavity (Fig. 3D) attracted digging at the base of wall-like cues (SI, Fig. S3). Note that such behavior can be interpreted as termites digging along the humidity gradient, i.e. toward the most humid region. Accordingly, in many preliminary experiments we observed that, in the absence of loose pellets, spontaneous digging usually started right above the hydration holes of our setup (Fig. 3B.)

In our study, we have outlined a general mechanism capable of organizing termite building activity: termites would focus pellet depositions at regions of strong evaporation flux. In turn, evaporation flux co-varies with surface curvature, which implies that the building rule is embedded in the shape itself of the nest internal structure.

One may wonder to what extent the simple building rule that we identify here generalizes to explain the nest-building behavior of larger termite colonies in the field, and whether the same building rules are shared across different termite species. The nests built by termites of different genera or species show a large diversity of forms (see e.g. Grassé, 1984), which indicates that the nest-building process should also be different. Arguably, the nest building behavior of termites, shaped by millions of years of evolution, must rely on more complex “building rules” than the simple ones highlighted here. Nonetheless, it is interesting to notice that the nests built by all species rely on a small number of architectural elements such as pillars and branching surfaces. We can imagine that, perhaps, simple variation of the basic building pattern described here, coupled with variation of the substrate evaporation itself (e.g. under the effect of air currents, the properties of the building material, and heat produced by the colony itself) would still be sufficient to explain a large part of termite nest diversity. Ocko and collaborators (2019) have already shown that a single mechanism can be responsible for determining the overall shape of nests made by various species: perhaps an equally simple general mechanism can account for the even larger variation of internal nest structure.

Added in the revised version: Beyond the field of termite architecture, many other biological structures, particularly at the microscopic scale, present convoluted structures that appear via tip growth and branching phenomena. Similar to termite nests, the morphogenesis of these structures may rely on the coupling between the local curvature of the growing substrate and the localization of cues that stimulate the growth. Generally, there is a growing consensus today on the fact that local curvature can affect cell migrations, cell patterning and tissue growth typically by direct mechanical sensing through the cytoskeleton (Pieuchot et al., 2018; Callens et al., 2023). However, analog mechanisms to those described here can also explain tip growth and branching in mammal lungs, Clément et al. (2012), and reconnections in the gastrovascular networks of jellyfishes Song et al. (2023), the common feature being that curvature guidance acts indirectly, by focusing the gradient of a diffusing quantity that promotes the substrate growth.

Methods and Materials

In our experiments, we monitor the building behavior of small experimental groups of Coptotermes gestroi termites confronted with a thin layer of clay and pre-prepared clay features. We image experimental trials for their entire duration and we analyze termite activity with custom made digital image processing routines. In parallel we run two types of control experiments without termites to obtain a non-intrusive estimation of temperature, humidity and evaporation field in our experimental setup. These experiments are described at the end of this section.

Experimental setup

The experimental setup is sketched in Fig. 1 (left) and can be described as follows. A fixed quantity (2.8 g) of red humid clay paste is flattened to form a disk (ø 5 cm) and placed in the center of a Petri dish (ø 8.5 cm). A system of 40 small holes (ø 0.8 mm) drilled in the bottom of the Petri dish keeps the clay paste hydrated sucking distilled water from a patch of wet cotton below the Petri dish. Two types of topographic cues molded in clay, can then be added at the center of the disk: 2 pillars 6 mm high and 8 mm apart or 1 wall 6 mm high and 12 mm long. The pillars are obtained pressing clay in a small eppendorf tube. The walls are obtained by smoothing a wedge of clay generated by rolling out a piece of clay in the dihedron between the table and the edge of a plastic ruler. Finally, in a circular band (1 cm large) halfway between the clay disk center and its edge, we add 0.12 g of sparse pellets of gray clay. To ensure the good size distribution, pellets are obtained from previous experiments and sterilized at 100 ?C for one hour to remove any possible chemical marker. Before the start of each experiment, a surface scan of the setup was taken using a NextEngine 3D Scanner ULTRA HD.

Termite colonies

Experimental groups were collected from a master captive colony of Coptotermes gestroi hosted at the LEEC laboratory (Villetaneuse, France) in a tropical room with constant temperature (26±2 °C) and relative humidity (70±10%), imitating their natural environment (see SI section S.I for a description of the termite species). Workers were attracted with humid towels and gently shoveled with a pencil on a plastic tray. Groups of 50 workers and 5 soldiers were then formed using an insect forceps and added to an experimental setup. While the procedure might be potentially stressful to termites, mortality was negligible throughout all the experiments. We ran 16 experiments with pillar cues, 11 with a wall and 6 with no cues, as summarized in table S1 (SI). An example of the structures spontaneously built by these termites within the plastic barrel which hosts the captive colony is shown in figure S1 (SI).

Recordings

A led lamp constantly lightened the setup from above. Top view pictures of the experimental setups were taken at regular intervals of 20 seconds during at least 24 hours using a reflex camera Nikon D7500. By subtracting the initial images and applying a median filter, we get rid of termites and the background as detailed in SI section S.III. Then, the color contrast between the clay disk and the pellets allowed identifying where pellets were collected (dark spots) and deposited (bright spots) and building the heatmap of both activities as a function of time. A subset of the experiments was recorded continuously using a 12 Mp usb-Camera (MER2-1220-32U3C) at 7fps. Videos were then analyzed using the open-source tracking tool Trex (Walter and Couzin, 2021) and the occupancy frequency of termites computed in each part of the experimental setup.

Temperature and humidity measurements

Temperature and humidity were measured using a commercial temperature-humidity probe (DHT22) connected to a Raspberry Pi. To not interfere with termites behavior, our temperature and humidity measurements were performed in a control experiment which was prepared using the same protocol as the others but where no termites were added. The probe was kept at two different positions that are i) at center of the experimental setup sitting on the clay disk and ii) at the periphery of the Petri dish, sitting on the bare plastic.

Growth model

Our growth model is the same described in Facchini et al. (2020) which consists in one single non-linear phase field equation: where the scalar f takes values between 0 and 1 and identifies the presence of the nest wall (f < 0.5) or the empty space (f > 0.5), and n = ∇f/|∇f| is the normal vector at the wall surface f = 0.5. One recognizes a growth term proportional to the surface curvature −∇ ⋅ n/2 which translates our main hypothesis on construction behavior, and a curvature diffusion term Δ∇⋅ n which mimics the smoothing behavior shown by termites (see Facchini et al., 2020, and citations inside) and the fact that there is a cutoff to the size of pellets added by termites. Finally, the non linear prefactor f (1−f) restricts the growth process to the wall surface, which is also coherent with termites behavior. Note that the simulations shown here are obtained approximating ∇ ⋅ n ≈ Δf as in Facchini et al. (2020). The parameter d selects the typical length scale of the expressed pattern and the cutoff scale below which features are smoothed out. At present, our model is still phenomenological and we cannot provide an estimation of d based on the experimental measurements of a limited set of physical quantities. Thus, we tune the parameter d to select a typical length scale which matches that of our topographic cues, that is the thickness of our clay pillars and walls (3 mm), which in turn was chosen to be of the same order of the termite size.

Our simulation are initialized using 3D copies of the experimental setups at time zero that are obtained as it follows. First, we obtain a surface scan of the experimental setup in the form of a 3D mesh using a surface 3D scanner (see section S.III in the Supplementary Information). Then, we interpolate the mesh on a 3D regular grid and assign the initial value of the scalar field f, setting f = 1 for the points that are below the clay surface and f = 0 for the points that are above the clay surface. Finally, a Gaussian filter is applied to unsharp the transition of f at the surface. Similarly to our previous publication, we also assume that the voxels where f > 0.85 at t = 0 cannot change in the following. This translates the fact that structures built by termites are not observed to be rearranged after they have dried and that in our experiments termites are prompted to collect pellets instead of digging.

Numerical simulations of the humidity field

In order to characterize the humidity field around the building substrate, we implemented simulations using the finite elements platform COMSOL Multiphysics. The simulations are based on the assumption (see section X.VI in the SI for a discussion) that humidity transport is diffusive within a cubic domain of side 18 mm whose bottom face is replaced by a 3D copy of the experimental setup. The boundary conditions on the humidity field h were no-flux on the lateral boundaries and Dirichlet on the bottom and top plate where the relative humidity h was fixed to 100% and 70%, to mimic the experimental value of h respectively at the clay disk surface and in the air (far from the disk surface) of our experimental room.

At a given temperature, relative humidity and absolute humidity are proportional, so assuming here that temperature is constant, we always refer to relative humidity and our simulations measure the humidity gradient at any position in units of mm−1 (exactly like curvature).

Saline solution experiments

We performed control experiments with no termites and hydrating the clay disk with a saline solution instead of distilled water to map the distribution of evaporation flux. Saline solution was prepared adding 8 g of NaCHO3to 100 ml of tap water.

Supporting information